<img style="float: left;" src="earth-lab-logo-rgb.png" width="150" height="150" />

# Homework Template: Earth Analytics Python Course: Spring 2020

Before submitting this assignment, be sure to restart the kernel and run all cells. To do this, pull down the Kernel drop down at the top of this notebook. Then select **restart and run all**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below.

* IMPORTANT: Before you submit your notebook, restart the kernel and run all! Your first cell in the notebook should be `[1]` and all cells should run in order! You will lose points if your notebook does not run. 

For all plots and code in general:

* Add appropriate titles to your plot that clearly and concisely describe what the plot shows (e.g. time, location, phenomenon).
* Be sure to use the correct bands for each plot.
* Specify the source of the data for each plot using a plot caption created with `ax.text()`.
* Place ONLY the code needed to create a plot in the plot cells. Place additional processing code ABOVE that cell (in a separate code cell).

Make sure that you:

* **Only include the package imports, code, data, and outputs that are CRUCIAL to your homework assignment.**
* Follow PEP 8 standards. Use the `pep8` tool in Jupyter Notebook to ensure proper formatting (however, note that it does not catch everything!).
* Keep comments concise and strategic. Don't comment every line!
* Organize your code in a way that makes it easy to follow. 
* Write your code so that it can be run on any operating system. This means that:
   1. the data should be downloaded in the notebook to ensure it's reproducible.
   2. all paths should be created dynamically using the os package to ensure that they work across operating systems. 
* Check for spelling errors in your text and code comments


```python
NAME = ""
COLLABORATORS = ""
```

![Colored Bar](colored-bar.png)

---
layout: single
category: courses
title: "Learn to Create Efficient Data Workflows in Python"
permalink: /courses/earth-analytics-python/create-efficient-data-workflows/
modified: '{:%Y-%m-%d}'.format(datetime.now())
week-landing: 10
week: 10
sidebar:
  nav:
comments: false
author_profile: false
course: "earth-analytics-python"
module-type: 'session'
---


{% include toc title="This Week" icon="file-text" %}

<div class="notice--info" markdown="1">

## <i class="fa fa-ship" aria-hidden="true"></i> Welcome to Week {{ page.week }}!

Welcome to week {{ page.week }} of Earth Analytics! This week you will learn how to automate a workflow using `Python`. You will design and implement your own workflow in `Python` that builds on the skills that you have learned in this course, such as functions and loops. You will also learn how to programmatically build paths to directories and files as well as parse strings to extract information from file and directory names.  

{% include/data_subsets/course_earth_analytics/_data-landsat-automation.md %}


</div>


## Automate a Workflow in Python

For this week’s assignment, you will generate a plot of the normalized difference vegetation index (NDVI) for two different locations in the United States to begin to understand how the growing seasons vary in each site:

1. <a href="https://www.neonscience.org/field-sites/field-sites-map/SJER" target="_blank">San Joaquin Experimental Range (SJER) in Southern California, United States</a>
2. <a href="https://www.neonscience.org/field-sites/field-sites-map/HARV" target="_blank">Harvard Forest (HARV) in the Eastern United States</a> 

From this plot, you will be able to compare the seasonal vegetation patterns of the two locations. This comparison would be useful if you were planning NEON’s upcoming flight season in both locations and wanted to ensure that you flew the area when the vegetation was the most green! If could also be useful if you wanted to track green-up as it happened over time in both sites to see if there were changes happening. 

As a bonus, you will also create a stacked NDVI output data product to share with your colleagues. You are doing all of the work to clean and process the data. It would be nice if you could share a data product output to save others the hassle. 

## Design A Workflow 

Your goal this week is to calculate the mean NDVI value for each Landsat 8 scene captured for a NEON site over a year. You have the following data to do accomplish this goal:

1. One year worth of Landsat 8 data for each site: Remember that for each landsat scene, you have a series of geotiff files representing bands and qa (quality assurance) layers in your data.
2. A site boundary “clip file” for each site: This is a shapefile representing the boundary of each NEON site. You will want to clip your landsat data to this boundary.

Before writing `Python` code, write pseudocode for your implementation. Pseudo-coding means that you will write out all of the steps that you need to perform. Then, you will identify areas where tasks are repeated that could benefit from a function, areas where loops might be appropriate, etc.  


## Homework for this Week

Your homework for this week is provided in the assignment dropbox on Canvas.


## Homework Plots

The plots below are examples of what your output plots will look like with and without dealing with clouds.


```python
# Autograding imports - do not modify this cell

import matplotcheck.autograde as ag
import matplotcheck.notebook as nb
import matplotcheck.timeseries as ts
from datetime import datetime
```

<!-- grade_id: cell-2d2391d865730d53, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1e463322f3d6c1a8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d86b87f27caf0d5e, points: 0, grade: true, solution: false -->
```python
# Test your dataframe! For each test, rename the dataframe comparison to check against your dataframe. 

# For example, in test one, the date test, change "student_dataframe['date_column']" to match
# your dataframe's name and date column. If your dataframe was called ndvi_data and the dates column was named
# date_time, than the lines should look like `assert date_series.equals(ndvi_data['date_time'].sort_values())`

# These tests are not graded so if they fail but your graph is correct
# you won't be docked any points! Leave the sort values statement, only replace the dataframe and column names.

# This is for data that hasn't been masked yet, it should be a half way sanity check. 

try:
    assert site_series.equals(ndvi_ts['site'].sort_values())
    print('Your site names appear to be accurate!')
except AssertionError:
    print('There seems to be an issue with your site names, check that column to make sure the values are correct.')
try:
    assert date_series.equals(ndvi_ts['date'].sort_values())
    print('Your date values appear to be accurate!')
except AssertionError:
    print('There seems to be an issue with your date values, check that column to make sure the values are correct.')
try:
    assert ndvi_series.equals(ndvi_ts['ndvi'].sort_values())
    print('Your ndvi values appear to be accurate!')
except AssertionError:
    print('There seems to be an issue with your ndvi values, check that column to make sure the values are correct.')
```

<!-- grade_id: cell-094ad248a67beb76, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f71de353817fae9e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0dbb343e769a09db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10287141af53f723, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

### DO NOT REMOVE LINE BELOW ###
final_masked_solution = nb.convert_axes(plt, which_axes="current")
```

<!-- grade_id: cell-963612e7cefa2a5e, points: 15, grade: true, solution: false -->
```python

```